Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Problem 1: Word Frequencies

In this problem, you will need to count the frequencies of all the words contained in a compressed text file supplied as input. We'll go through this process in several of parts.

## Part A (2 points)

Write a function which takes in a string and and returns the frequencies of each word in a dictionary. Consider 'words' simply the result of splitting the input text into a list, using any form of whitespace as a separator. This is obviously a very naive definition of 'word', but it shall suffice for the purposes of this exercise.

<!-- grade_id: cell-d61445e6794c93bf, points: 0, grade: false, solution: true -->
```python
def word_freq(text):
    """Return a dictionary of word frequencies for the given text."""

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2a080a2b8fd3688f, points: 2, grade: true, solution: false -->
```python
"""Check that word_freq returns the right values"""

text = "do we see how many times do we see many very many"
freqs = word_freq(text)
assert freqs["many"] == 3
assert freqs["do"] == 2
assert freqs["how"] == 1
```

## Part B (3 points)
Using your previous function, write a function which takes in a string and returns a list of (word, frequency) tuples in order of frequency (decreasing)

<!-- grade_id: cell-90337e2c390fe07e, points: 0, grade: false, solution: true -->
```python
def ordered_kv(text):
    """Return a list of (word, frequency) tuples ordered by decreasing freq"""
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2e70179247e0d442, points: 3, grade: true, solution: false -->
```python
"""Check that ordered_kv returns the right values"""

text = "do we see how many times do we see many very many"
kv = ordered_kv(text)
assert kv[0][1] == 3
assert kv[0][0] == "many"
assert kv[-1][1] == 1
```

## Part C (3 points)
Using `ordered_kv`, write a function which takes in a path to a Gzipped file and returns a list of the top 20 words and their frequencies in decreasing order and list of the bottom 20 words and their frequencies in increasing order.

The output should look like:

```([("most_freqent_word", length_of_word), ("second_most_frequent_word", length_of_word),...], [("least_frequent_word", length_of_word), ("second_least_frequent_word", length_of_word),...])```

Note: This file is compressed!

<!-- grade_id: cell-eeebfc94eaaf7c9f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

def top_and_bottom(zipped_path):
    """Return the top 20 and bottom 20 words in a zipped file"""
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-eb209267c4369fb2, points: 1, grade: true, solution: false -->
```python
"""Check that top_and_bottom returns two lists"""
out = top_and_bottom('/accounts/class/s159/exchange/HISTORY.gz')
assert len(out) == 2
assert type(out[0]) == list
assert type(out[1]) == list
```

<!-- grade_id: cell-52aa059e022fb56b, points: 2, grade: true, solution: false -->
```python
"""Check that top_and_bottom returns the right lists"""
top20, bot20 = top_and_bottom('/accounts/class/s159/exchange/HISTORY.gz')

assert top20[0][0] == b"the"
assert top20[0][1] == 2461
assert top20[10][0] == b"The"
assert top20[10][1] == 581

for bot_word in bot20:
    assert bot_word[1] == 1
```

## Part D (2 points)
Write a function which takes in a path to a Gzipped file and prints the top 20 and bottom 20 words in the same order as above. Format the strings so that the function prints out:

```
The top 20 words are:
Word: word1, Frequency: freq1
...

The bottom 20 words are (word, frequency):
word1,freq1
...
```

Where `word1,freq1` are replaced with the correct words and frequencies 

<!-- grade_id: cell-b151d8fd8f60b0f5, points: 2, grade: true, solution: true -->
```python
def print_top_and_bottom(zipped_path):
    """Print the top 20 and bottom 20 words in a zipped file"""
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
# print output here
print_top_and_bottom('/accounts/class/s159/exchange/HISTORY.gz')
```

```python

```